---
title       : Giới thiệu cơ bản về Python
description : Trong chương này chúng ta sẽ được giới thiệu về Python cho Data Analysis, lý do nên sử dụng Python cho Data Analysis. Giới thiệu tổng quan nhất về hệ sinh thái của Python.
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:MultipleChoiceExercise lang:python xp:50 skills:1 key:6f76b6229e

## Why learn Python for data analysis

Python (an interpreted language) has gathered a lot of interest recently as a preferred choice of language for data analysis. Here are some reasons in favour of learning Python:

* It is open source – free to install and use
* Python has an awesome online community - latest algorithms come to Python in a matter of days
* It is easy to learn
* It can become a common language for data science and production of web-based analytics products

Which of the following is not a reason to learn Python for data analysis?

*** =instructions

* Python is easy to learn.
* Python is an interpreted language, so computation times can be higher than compiler based languages in some cases.
* Python has good libraries for data science.
* It is a production ready language (from web & software perspective).

*** =hint
Interpreted languages are typically easier to learn, but take longer computational time than compiler based languages.

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:132f57517d
## Python 2.7 vs. Python 3.5?

You will come across this question soon after you start using Python. Python has 2 popular competing versions. Both versions have their pros and cons.

**Benefits of Python 2.7**

* Awesome online community. Easier to find answers when you get stuck at places.
* Tonnes of third party libraries

**Benefits of Python 3.5**

* Cleaner and faster
* It is the future!



*** =instructions
Which version of Python would you recommend to someone who needs to use several third party libraries?

* Python 2.7
* Python 3.5
* Should work on both


*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```
